<h1 style="text-align:center; color:red;">Demo Assignment</h1>

<h3 style="color:red">Demonstration of ADD Function</h3>

<!-- grade_id: cell-f8f2704b389591e1, points: 0, grade: false, solution: true -->
```python
# Cell 1, cell metadata=autograded-answer
def add(a, b):
    ### BEGIN SOLUTION
    c = a + b
    ### END SOLUTION
    return c
```

<!-- grade_id: cell-6fdff3cccf499fea, points: 3, grade: true, solution: false -->
```python
# Cell 2, cell metadata=autograder-tests
assert add(1, 2) == 3
### BEGIN HIDDEN TESTS
assert add(-1, 2) == 1
### END HIDDEN TESTS
```

<h3 style="color:red">Demonstration of SUBTRACT Function</h3>

<!-- grade_id: cell-9d53fcc21662701a, points: 0, grade: false, solution: true -->
```python
# Cell 1, cell metadata=autograded-answer
def subtract(a, b):
    ### BEGIN SOLUTION
    c = a - b
    ### END SOLUTION
    return c
```

<!-- grade_id: cell-22d5407ed9757a16, points: 3, grade: true, solution: false -->
```python
# Cell 2, cell metadata=autograder-tests
assert subtract(1, 2) == -1
### BEGIN HIDDEN TESTS
assert subtract(-1, 2) == -3
### END HIDDEN TESTS
```

<h3 style="color:red">Demonstration of MULTIPLY Function</h3>

<!-- grade_id: cell-9b3bf79d6f7c753a, points: 0, grade: false, solution: true -->
```python
# Cell 1, cell metadata=autograded-answer
def multiply(a, b):
    ### BEGIN SOLUTION
    c = a * b
    ### END SOLUTION
    return c
```

<!-- grade_id: cell-fb83a713eb11ac1a, points: 5, grade: true, solution: false -->
```python
# Cell 2, cell metadata=autograder-tests
assert multiply(1, 2) == 2
### BEGIN HIDDEN TESTS
assert multiply(-1, 2) == -2
### END HIDDEN TESTS
```

<h3 style="color:red">Demonstration of DIVIDE Function</h3>

<!-- grade_id: cell-51c5c3b04bd1e4c9, points: 0, grade: false, solution: true -->
```python
# Cell 1, cell metadata=autograded-answer
def divide(a, b):
    ### BEGIN SOLUTION
    c = a / b
    ### END SOLUTION
    return c
```

<!-- grade_id: cell-18c539aa81d75695, points: 5, grade: true, solution: false -->
```python
# Cell 2, cell metadata=autograder-tests
assert divide(1, 2) == 0.5
### BEGIN HIDDEN TESTS
assert divide(-1, 2) == -0.5
### END HIDDEN TESTS
```